<a href="https://colab.research.google.com/github/Miriam-Ivy/Microsoft-Movies-Project/blob/master/Miriam_Checkpoint.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Phase 2 Code Challenge

This code challenge is designed to test your understanding of the Phase 2 material. It covers:

- Normal Distribution
- Statistical Tests
- Bayesian Statistics
- Linear Regression

_Read the instructions carefully_. You will be asked both to write code and to answer short answer questions.

## Code Tests

We have provided some code tests for you to run to check that your work meets the item specifications. Passing these tests does not necessarily mean that you have gotten the item correct - there are additional hidden tests. However, if any of the tests do not pass, this tells you that your code is incorrect and needs changes to meet the specification. To determine what the issue is, read the comments in the code test cells, the error message you receive, and the item instructions.

## Short Answer Questions

For the short answer questions...

* _Use your own words_. It is OK to refer to outside resources when crafting your response, but _do not copy text from another source_.

* _Communicate clearly_. We are not grading your writing skills, but you can only receive full credit if your teacher is able to fully understand your response.

* _Be concise_. You should be able to answer most short answer questions in a sentence or two. Writing unnecessarily long answers increases the risk of you being unclear or saying something incorrect.

```python
# Run this cell without changes to import the necessary libraries

import itertools
import numpy as np
import pandas as pd
from numbers import Number
from scipy import stats
import matplotlib.pyplot as plt
import seaborn as sns
%matplotlib inline

import pickle
```

---
## Part 1: Normal Distribution [Suggested time: 20 minutes]
---
In this part, you will analyze check totals at a TexMex restaurant. We know that the population distribution of check totals for the TexMex restaurant is normally distributed with a mean of \\$20 and a standard deviation of \\$3.

### 1.1) Create a numeric variable `z_score_26` containing the z-score for a \\$26 check.

**Starter Code**

    z_score_26 = 2

<!-- grade_id: cell-b947773fd0a27e5d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# Import necessary libraries
import numpy as np

mean = 20
std_dev = 3
check_value = 26

z_score_26 = (check_value - mean) / std_dev

print("Z-score for a $26 check:", z_score_26)


```

Output:
```
Z-score for a $26 check: 2.0
```

<!-- grade_id: cell-a72f04fffcbada78, points: 1, grade: true, solution: false -->
```python
# This test confirms that you have created a numeric variable named z_score_26

assert isinstance(z_score_26, Number)

```

### 1.2) Create a numeric variable `p_under_26` containing the approximate proportion of all checks that are less than \\$26.

Hint: Use the answer from the previous question along with the empirical rule, a Python function, or this [z-table](https://www.math.arizona.edu/~rsims/ma464/standardnormaltable.pdf).

**Starter Code**

    p_under_26 = 0.978

<!-- grade_id: cell-a223448d6f75a998, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE

from scipy import stats

mean = 20
std_dev = 3
check_value = 26

p_under_26 = norm.cdf(check_value, loc=mean, scale=std_dev)

print("Approximate proportion of checks less than $26:", p_under_26)


```

Output:
```
Approximate proportion of checks less than $26: 0.9772498680518208
```

<!-- grade_id: cell-626d2e6dea5a86ac, points: 1, grade: true, solution: false -->
```python
# This test confirms that you have created a numeric variable named p_under_26

assert isinstance(p_under_26, Number)

# These tests confirm that p_under_26 is a value between 0 and 1

assert p_under_26 >= 0
assert p_under_26 <= 1

```

### 1.3) Create numeric variables `conf_low` and `conf_high` containing the lower and upper bounds (respectively) of a 95% confidence interval for the mean of one waiter's check amounts using the information below.

One week, a waiter gets 100 checks with a mean of \\$19 and a standard deviation of \\$3.

**Starter Code**

    n = 100
    mean = 19
    std = 3
    
    conf_low = 18.412
    conf_high = 19.588

<!-- grade_id: cell-fb0a963f092cb64d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
import numpy as np
from scipy import stats

# Given values
sample_mean = 19
sample_std_dev = 3
sample_size = 100
confidence_level = 0.95

z_score = norm.ppf((1 + confidence_level) / 2)

standard_error = sample_std_dev / np.sqrt(sample_size)

margin_of_error = z_score * standard_error

conf_low = sample_mean - margin_of_error
conf_high = sample_mean + margin_of_error

print("Lower bound of the 95% confidence interval:", conf_low)
print("Upper bound of the 95% confidence interval:", conf_high)




```

Output:
```
Lower bound of the 95% confidence interval: 18.412010804637983
Upper bound of the 95% confidence interval: 19.587989195362017
```

<!-- grade_id: cell-d995ba92d5f4f1f3, points: 1, grade: true, solution: false -->
```python
# These tests confirm that you have created numeric variables named conf_low and conf_high

assert isinstance(conf_low, Number)
assert isinstance(conf_high, Number)

# This test confirms that conf_low is below conf_high

assert conf_low < conf_high

# These statements print your answers for reference to help answer the next question

print('The lower bound of the 95% confidence interval is {}'.format(conf_low))
print('The upper bound of the 95% confidence interval is {}'.format(conf_high))

```

Output:
```
The lower bound of the 95% confidence interval is 18.412010804637983
The upper bound of the 95% confidence interval is 19.587989195362017
```

### 1.4) Short Answer: Interpret the 95% confidence interval you just calculated in Question 1.3.

<!-- grade_id: cell-8212712ac7485fa9, points: 1, grade: true, solution: true -->
The 95% confidence interval suggests that we are 95% confident that the true mean of the waiter's check amounts falls within the range of $ -19 and  $ +19 margin of error.








---
## Part 2: Statistical Testing [Suggested time: 20 minutes]
---
The TexMex restaurant recently introduced queso to its menu.

We have a random sample containing 2000 check totals, all from different customers: 1000 check totals for orders without queso ("no queso") and 1000 check totals for orders with queso ("queso").

In the cell below, we load the sample data for you into the arrays `no_queso` and `queso` for the "no queso" and "queso" order check totals, respectively.

```python
# Run this cell without changes

# Load the sample data
no_queso = pickle.load(open('/content/queso.pkl', 'rb'))
queso = pickle.load(open('/content/queso.pkl', 'rb'))
```

### 2.1) Short Answer: State null and alternative hypotheses to use for testing whether customers who order queso spend different amounts of money from customers who do not order queso.

<!-- grade_id: cell-1992605f4c2b8751, points: 1, grade: true, solution: true -->
We use a two-tailed hypothesis test using the Null and Alternative hypothesis to determine whether there is a significant difference in the average check totals between customers who order queso and those who do not.








### 2.2) Short Answer: What would it mean to make a Type I error for this specific hypothesis test?

Your answer should be _specific to this context,_  not a general statement of what Type I error is.

<!-- grade_id: cell-f5e449677777b2fb, points: 1, grade: true, solution: true -->
A Type I error in this test would mean that there is a significant difference in the average check totals between customers who order queso and those who do not when, although this cannot happen in reality. It creates a false positive conclusion.

### 2.3) Create a numeric variable `p_value` containing the p-value associated with a statistical test of your hypotheses.

You must identify and implement the correct statistical test for this scenario. You can assume the two samples have equal variances.

Hint: Use `scipy.stats` to calculate the answer - it has already been imported as `stats`. Relevant documentation can be found [here](https://docs.scipy.org/doc/scipy/reference/stats.html#statistical-tests).

**Starter Code**

    p_value =

<!-- grade_id: cell-c2e90acfaebc8303, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2364f38312cbc15e, points: 1, grade: true, solution: false -->
```python
# These tests confirm that you have created a numeric variable named p_value

assert isinstance(p_value, Number)

```

### 2.4) Short Answer: Can you reject the null hypothesis using a significance level of $\alpha$ = 0.05? Explain why or why not.

<!-- grade_id: cell-92ea8835a9a326eb, points: 1, grade: true, solution: true -->
YOUR ANSWER HERE

---
## Part 3: Bayesian Statistics [Suggested time: 15 minutes]
---
A medical test is designed to diagnose a certain disease. The test has a false positive rate of 10%, meaning that 10% of people without the disease will get a positive test result. The test has a false negative rate of 2%, meaning that 2% of people with the disease will get a negative result. Only 1% of the population has this disease.

### 3.1) Create a numeric variable `p_pos_test` containing the probability of a person receiving a positive test result.

Assume that the person being tested is randomly selected from the broader population.

**Starter Code**
    
    false_pos_rate = 0.1
    false_neg_rate = 0.02
    population_rate = 0.01
    
    p_pos_test = 0.1088

<!-- grade_id: cell-30ff893c7c30f29d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
false_positive_rate = 0.1
false_negative_rate = 0.02
population_rate = 0.01
no_disease_prob = 1 - population_rate

p_pos_test = (1 - false_negative_rate) * population_rate + false_positive_rate * no_disease_prob

print("Probability of a positive test result:", p_pos_test)


```

Output:
```
Probability of a positive test result: 0.10880000000000001
```

<!-- grade_id: cell-2d49a1068c9e46b2, points: 1, grade: true, solution: false -->
```python
# This test confirms that you have created a numeric variable named p_pos_test

assert isinstance(p_pos_test, Number)

# These tests confirm that p_pos_test is a value between 0 and 1

assert p_pos_test >= 0
assert p_pos_test <= 1

```

### 3.2) Create a numeric variable `p_pos_pos` containing the probability of a person actually having the disease if they receive a positive test result.

Assume that the person being tested is randomly selected from the broader population.

Hint: Use your answer to the previous question to help answer this one.

**Starter Code**
    
    false_pos_rate = 0.1
    false_neg_rate = 0.02
    population_rate = 0.01
    
    p_pos_pos = 0.09

<!-- grade_id: cell-742273642f14d1ab, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
# Given values
false_negative_rate = 0.02
population_rate = 0.01

p_positive_given_disease = 1 - false_negative_rate

p_pos_test = 0.1088

p_disease_given_pos_test = (p_positive_given_disease * population_rate) / p_pos_test

p_pos_pos = p_disease_given_pos_test

print("Probability of having the disease given a positive test result:", p_pos_pos)



```

Output:
```
Probability of having the disease given a positive test result: 0.0900735294117647
```

<!-- grade_id: cell-4825f55ab11b5fff, points: 1, grade: true, solution: false -->
```python
# This test confirms that you have created a numeric variable named p_pos_pos

assert isinstance(p_pos_pos, Number)

# These tests confirm that p_pos_pos is a value between 0 and 1

assert p_pos_pos >= 0
assert p_pos_pos <= 1

```

---

## Part 4: Linear Regression [Suggested Time: 20 min]
---

In this section, you'll run regression models with advertising data. The dataset includes the advertising spending and sales for 200 products. These products are a random sample of 200 products from a larger population of products.

The dataset has three features - `TV`, `radio`, and `newspaper` - that describe how many thousands of advertising dollars were spent promoting the product. The target, `sales`, describes how many millions of dollars in sales the product had.

```python
import pandas as pd

data = pd.read_csv('/advertising.csv')

```

```python
# Run this cell without changes

import statsmodels
import statsmodels.api as sm

data = pd.read_csv('/advertising.csv').drop('Unnamed: 0', axis=1)
X = data.drop('sales', axis=1)
y = data['sales']
```

### 4.1) Create a variable `tv_mod_summ` containing the summary of a fit StatsModels `ols` model for a linear regression using `TV` to predict `sales`.

**Starter Code**

    tv_mod = ols(
    tv_mod_summ = tv_mod.summary()

<!-- grade_id: cell-86323454c9b2aa75, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
import pandas as pd
from statsmodels.formula.api import ols

data = pd.read_csv('./advertising.csv').drop('Unnamed: 0', axis=1)

tv_mod = ols('sales ~ TV', data=data).fit()

tv_mod_summ = tv_mod.summary()

print(tv_mod_summ)


```

Output:
```
OLS Regression Results                            
==============================================================================
Dep. Variable:                  sales   R-squared:                       0.612
Model:                            OLS   Adj. R-squared:                  0.610
Method:                 Least Squares   F-statistic:                     312.1
Date:                Sun, 17 Dec 2023   Prob (F-statistic):           1.47e-42
Time:                        19:10:36   Log-Likelihood:                -519.05
No. Observations:                 200   AIC:                             1042.
Df Residuals:                     198   BIC:                             1049.
Df Model:                           1                                         
Covariance Type:            nonrobust                                         
==============================================================================
                 coef    std err          t      P>|t|      [0.025      0.975]
-
```

<!-- grade_id: cell-84e163360346b9f6, points: 1, grade: true, solution: false -->
```python
# This test confirms that you have created a variable named tv_mod_summ containing a StatsModels Summary

assert type(tv_mod_summ) == statsmodels.iolib.summary.Summary

# This statement prints your answer for reference to help answer the next question

print(tv_mod_summ)

```

Output:
```
OLS Regression Results                            
==============================================================================
Dep. Variable:                  sales   R-squared:                       0.612
Model:                            OLS   Adj. R-squared:                  0.610
Method:                 Least Squares   F-statistic:                     312.1
Date:                Sun, 17 Dec 2023   Prob (F-statistic):           1.47e-42
Time:                        19:10:36   Log-Likelihood:                -519.05
No. Observations:                 200   AIC:                             1042.
Df Residuals:                     198   BIC:                             1049.
Df Model:                           1                                         
Covariance Type:            nonrobust                                         
==============================================================================
                 coef    std err          t      P>|t|      [0.025      0.975]
-
```

### 4.2) Short Answer: Is there a statistically significant relationship between TV advertising spend and sales in this model? How did you determine this from the model output?

This question is asking you to use your findings from the sample in your dataset to make an inference about the relationship between TV advertising spend and sales in the broader population.

<!-- grade_id: cell-ebd4cb97464e0212, points: 1, grade: true, solution: true -->
Yes. This is because the significant relationship is determined by the 'p-value'in relationship with the 'TV'variable. When the 'p-value'is less than a 0.05 significance level, then there is a close relationship.

### 4.3) Short Answer: Run the cell below to produce a correlation matrix. Given the output, would you expect any collinearity issues if you included all of these features in one regression model?

```python
# Run this cell

X.corr()
```

Output:
```
TV     radio  newspaper
TV         1.000000  0.054809   0.056648
radio      0.054809  1.000000   0.354104
newspaper  0.056648  0.354104   1.000000
```

<!-- grade_id: cell-c1e98a26b16748b1, points: 1, grade: true, solution: true -->
no. If you observe high absolute values (close to 1) in the correlation matrix between pairs of independent variables, it suggests potential collinearity. Thus we can assume there is NO collinearity.

### 4.4) Create a variable `all_mod_summ` containing the summary of a fit StatsModels `ols` model for a multiple regression using `TV`, `radio`, and `newspaper` to predict `sales`.

**Starter Code**

    all_mod = ols(
    all_mod_summ = all_mod.summary()

<!-- grade_id: cell-a9cc84d7d7666b0f, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
import pandas as pd
from statsmodels.formula.api import ols

formula = 'sales ~ TV + radio + newspaper'

all_mod = ols(formula, data=data).fit()

all_mod_summ = all_mod.summary()

print(all_mod_summ)

correlation_matrix = data[['TV', 'radio', 'newspaper']].corr()

print("\nCorrelation Matrix:")
print(correlation_matrix)

threshold = 0.8

highly_correlated_pairs = []

# Iterate through the correlation matrix
for i in range(len(correlation_matrix.columns)):
    for j in range(i + 1, len(correlation_matrix.columns)):
        if abs(correlation_matrix.iloc[i, j]) > threshold:
            pair = (correlation_matrix.columns[i], correlation_matrix.columns[j])
            highly_correlated_pairs.append(pair)

# Display highly correlated pairs
if len(highly_correlated_pairs) > 0:
    print("\nHighly Correlated Variable Pairs:")
    for pair in highly_correlated_pairs:
        print(f"{pair[0]} and {pair[1]}")
else:
    print("\nNo highly correlated variable pairs found.")
```

Output:
```
OLS Regression Results                            
==============================================================================
Dep. Variable:                  sales   R-squared:                       0.897
Model:                            OLS   Adj. R-squared:                  0.896
Method:                 Least Squares   F-statistic:                     570.3
Date:                Sun, 17 Dec 2023   Prob (F-statistic):           1.58e-96
Time:                        20:17:50   Log-Likelihood:                -386.18
No. Observations:                 200   AIC:                             780.4
Df Residuals:                     196   BIC:                             793.6
Df Model:                           3                                         
Covariance Type:            nonrobust                                         
==============================================================================
                 coef    std err          t      P>|t|      [0.025      0.975]
-
```

<!-- grade_id: cell-1f7f745c780b07b6, points: 1, grade: true, solution: false -->
```python
# This test confirms that you have created a variable named all_mod_summ containing a StatsModels Summary

assert type(all_mod_summ) == statsmodels.iolib.summary.Summary

# This statement prints your answer for reference to help answer the next question

print(all_mod_summ)

```

Output:
```
OLS Regression Results                            
==============================================================================
Dep. Variable:                  sales   R-squared:                       0.897
Model:                            OLS   Adj. R-squared:                  0.896
Method:                 Least Squares   F-statistic:                     570.3
Date:                Sun, 17 Dec 2023   Prob (F-statistic):           1.58e-96
Time:                        20:04:16   Log-Likelihood:                -386.18
No. Observations:                 200   AIC:                             780.4
Df Residuals:                     196   BIC:                             793.6
Df Model:                           3                                         
Covariance Type:            nonrobust                                         
==============================================================================
                 coef    std err          t      P>|t|      [0.025      0.975]
-
```

### 4.5) Short Answer: Does this model do a better job of explaining sales than the previous model using only the `TV` feature? Explain how you determined this based on the model output.

<!-- grade_id: cell-7e27f10fcaaed6a9, points: 1, grade: true, solution: true -->
Yes. we are given the information in the notes.